<img src='http://www.u.arizona.edu/~hammond/hltlogo1.png' style="float:right">

Linguistics 578<br>
Fall 2022<br>
Hammond

## Things to remember about any homework assignment:

1. For this assignment, you will edit this jupyter notebook and turn it in. Do not turn in pdf files or separate `.py` files.
1. Late work is not accepted.
1. Given the way I grade, you should try to answer *every* question, even if you don't like your answer or have to guess.
1. You may *not* use `python` modules that we have not already used in class.
1. You may certainly talk to your classmates about the assignment, but everybody must turn in *their own* work. It is not acceptable to turn in work that is essentially the same as the work of classmates.
1. All code must run. It doesn't have to be perfect, it may not do all that you want it to do, but it must run without error.
1. Code must run in reasonable time. Assume that if it takes more than *5 minutes* to run (on your machine), that's too long.
1. Please do not add, remove, or copy autograded cells.
1. Make sure to select `restart, run all cells` from the `kernel` menu when you're done and before you turn this in!

***

***my name***: Rachel Hansen

***people I talked to about the assignment***: My cat, Freya, she's an expert at sine waves

***

## Homework #1

Here are the imports. Please do not import anything else.

```python
import numpy as np
from scipy.fftpack import rfft
```

1. If we calculate sine from $0$ to $40$, how many cycles is that? (Be precise to two decimal points.)

<!-- grade_id: q1a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
cycles = 40/(2*np.pi)
x="{:.2f}".format(cycles)
#raise NotImplementedError()
#x = ?
print (x)
```

Output:
```
6.37
```

<!-- grade_id: q1t1, points: 1, grade: true, solution: false -->
```python
#this test is hidden!
```

2. Write a function that will return a sine wave with `freq` cycles over `samples` samples.

<!-- grade_id: q2a, points: 0, grade: false, solution: true -->
```python
def makesine(freq,samples):
    '''make a sine wave
    args:
        freq: how many cycles
        samples: how many samples
    returns:
        a numpy array with the sine wave
    '''
    numberoftimes = samples/(2*np.pi)
    v = freq/numberoftimes
    x = np.arange(0,samples, dtype=int)
    print (len(np.sin(x * v)))
    return(np.sin(x * v))

 #   raise NotImplementedError()
```

<!-- grade_id: q2t1, points: 1, grade: true, solution: false -->
```python
res = makesine(10,1000)
assert len(res) == 1000
```

Output:
```
1000
```

<!-- grade_id: q2t2, points: 1, grade: true, solution: false -->
```python
assert np.isclose(max(res),1,atol=.1)
```

<!-- grade_id: q2t3, points: 1, grade: true, solution: false -->
```python
assert np.isclose(min(res),-1,atol=.1)
```

<!-- grade_id: q2t4, points: 1, grade: true, solution: false -->
```python
zeros = 0
for i in range(1,len(res)):
    if res[i]*res[i-1] <= 0: zeros += 1
zeros //= 2
assert zeros == 10
```

3. Now write a function that will return a sine wave with a specific frequency, sample rate, *and* duration (in seconds). (You may call your previous function or not, as you choose.)

<!-- grade_id: q3a, points: 0, grade: false, solution: true -->
```python
def makewave(freq,samples,duration):
    '''make a sine wave with a specific duration
    args:
        freq: frequency of the wave
        samples: how many samples per second
        duration: how long (in seconds)
    returns:
        the wave as a numpy array
    '''
    samples = samples * duration
    freq = freq * duration
    numberoftimes = samples/(2*np.pi)
    v = freq/numberoftimes
    x = np.arange(0,samples)
  #  print (len(x))
  #  print(len(np.sin(x)))
    return(np.sin(x * v))
  #  raise NotImplementedError()
```

```python
makewave(100,1000,1.2)
```

Output:
```
array([ 0.        ,  0.58778525,  0.95105652, ..., -0.95105652,
       -0.95105652, -0.58778525])
```

<!-- grade_id: q3t1, points: 1, grade: true, solution: false -->
```python
res2 = makewave(100,1000,1.2)
assert len(res2) == 1200
```

<!-- grade_id: q3t2, points: 1, grade: true, solution: false -->
```python
assert np.isclose(max(res2),1,atol=.1)
```

<!-- grade_id: q3t3, points: 1, grade: true, solution: false -->
```python
assert np.isclose(min(res2),-1,atol=.1)
```

<!-- grade_id: q3t4, points: 1, grade: true, solution: false -->
```python
zeros = 0
for i in range(1,len(res2)):
    if res2[i]*res2[i-1] <= 0: zeros += 1
zeros //= 2
print (zeros)
assert zeros == 120
```

Output:
```
120
```

4. Now write a function that returns whether some hypothetical harmonic is consistent with a specific fundamental frequency.

<!-- grade_id: q4a, points: 0, grade: false, solution: true -->
```python
def harmonic(f0,h):
    '''checks if a harmonic frequency is legit
    args:
        f0: fundamental frequency of the complex wave
        h: harmonic frequency
    returns:
        boolean
    '''
    # YOUR CODE HERE
    return h % f0 == 0

    raise NotImplementedError()
```

<!-- grade_id: q4t1, points: 1, grade: true, solution: false -->
```python
assert harmonic(50,150)
```

<!-- grade_id: q4t2, points: 1, grade: true, solution: false -->
```python
assert harmonic(50,50)
```

<!-- grade_id: q4t3, points: 1, grade: true, solution: false -->
```python
assert harmonic(60,100) == False
```

5. Write a function that will generate a pulse train. Your function will take four arguments: the fundamental frequency, the maximum frequency for the harmonics, the duration (in seconds), and the sample rate (per second). All harmonics should have the same amplitude and the maximum amplitude of the complex wave should be $1$. (You may call your previous functions if you choose.)

<!-- grade_id: q5a, points: 0, grade: false, solution: true -->
```python
def pulsetrain(f0,freqmax,duration,samples):
    '''make a pulse train
    args:
        f0: fundamental frequency of the complex wave
        freqmax: maximum frequency for harmonics
        duration: duration in seconds
        samples: sample rate (per second)
    returns:
        wave as numpy array
    '''
    # YOUR CODE HERE
    freqarray = []
    tempfreq = f0
    while tempfreq <= freqmax:
      freqarray.append(tempfreq)
      tempfreq = tempfreq + f0
    
    wavesarray = []
    for freq in freqarray:
      wave = makewave(freq, samples, duration).tolist()
      wavesarray.append(wave)

    pulsearray = []
    for i in range(len(wavesarray[0])):
      pulsearray.append(0)
      for wave in wavesarray:
        pulsearray[-1] = pulsearray[-1] + wave[i]
    
    for i in range(len(pulsearray)):
      if pulsearray[i] > 1.0:
        pulsearray[i] = 1.0
      elif pulsearray[i] < -1.0:
        pulsearray[i] = -1.0

    return np.array(pulsearray)  
  # raise NotImplementedError()
```

<!-- grade_id: q5t1, points: 1, grade: true, solution: false -->
```python
res3 = pulsetrain(100,2000,1,10000)
assert len(res3) == 10000
```

<!-- grade_id: q5t2, points: 1, grade: true, solution: false -->
```python
assert np.isclose(res3.max(),1.0,atol=.1)
```

<!-- grade_id: q5t3, points: 1, grade: true, solution: false -->
```python
yf = rfft(res3)
assert yf[800] < -250
```

<!-- grade_id: q5t4, points: 1, grade: true, solution: false -->
```python
assert yf[3800] < -250
```